---
title: Chapter Title Here
description: Chapter description goes here.
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills:
  - 2
```

Welcome o Nuno's first exercise.

Mar 18, 11:34am

`@instructions`
This is a really easy exercise, and you should be able to do it without any previous knowledge (only common sense).

To solve this exercise, just write the code to print the "Hello Nuno" string.

Good luck!

`@hint`
Use common sense to solve this exercise

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
# use print function to print 'Hello Nuno' string
print('type something here')
```

`@solution`
```{python}
print('Hello Nuno')
```

`@sct`
```{python}
Ex().has_output('Hello Nuno')
```

---

## From GitHub

```yaml
type: VideoExercise
key: d47aeb2160
xp: 50
```

`@projector_key`
f87f31298d11674297bff2f85c0eb420

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: f03997e087
xp: 50
```

<!-- Guidelines for the question: https://instructor-support.datacamp.com/en/articles/2375516-course-multiple-choice-exercises. -->

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@possible_answers`
- [Correct answer 1]
- Wrong answer 2
- Wrong answer 3

`@feedback`
<!-- Examples of good feedback messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.  -->
- Perfect!
- Error message answer 2
- Error message answer 3
